<div class="alert alert-info">
    
➡️ Before you start, make sure that you are familiar with the **[study guide](https://liu-nlp.ai/text-mining/logistics/)**, in particular the rules around **cheating and plagiarism** (found in the course memo).

➡️ If you use code from external sources (e.g. StackOverflow, ChatGPT, ...) as part of your solutions, don't forget to add a reference to these source(s) (for example as a comment above your code).

➡️ Make sure you fill in all cells that say **`YOUR CODE HERE`** or **YOUR ANSWER HERE**.  You normally shouldn't need to modify any of the other cells.

</div>

# L1: Information Retrieval

In this lab you will apply basic techniques from information retrieval to implement the core of a minimalistic search engine. The data for this lab consists of a collection of app descriptions scraped from the [Google Play Store](https://play.google.com/store/apps?hl=en). From this collection, your search engine should retrieve those apps whose descriptions best match a given query under the vector space model.

```python
# Define some helper functions that are used in this notebook
from IPython.display import display, HTML

def success():
    display(HTML('<div class="alert alert-success"><strong>Checks have passed!</strong></div>'))
```

## Dataset

The app descriptions come in the form of a compressed [JSON](https://en.wikipedia.org/wiki/JSON) file. Start by loading this file into a [Pandas](https://pandas.pydata.org) [DataFrame](https://pandas.pydata.org/pandas-docs/stable/getting_started/dsintro.html#dataframe).

```python
import bz2
import numpy as np
import pandas as pd
pd.set_option('display.max_colwidth', 500)

with bz2.open('app-descriptions.json.bz2', mode='rt', encoding='utf-8') as source:
    df = pd.read_json(source, encoding='utf-8')
```

In Pandas, a DataFrame is a table with indexed rows and labelled columns of potentially different types. You can access data in a DataFrame in various ways, including by row and column. To give an example, the code in the next cell shows rows 200–204:

```python
df.loc[200:205]
```

Output:
```
name  \
200  Brick Breaker Star: Space King   
201      Brick Classic - Brick Game   
202     Bricks Breaker - Glow Balls   
203            Bricks Breaker Quest   
204             Brothers in Arms® 3   
205       Brown Dust - Tactical RPG   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                             description  
200  Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n-
```

As you can see, there are two labelled columns: `name` (the name of the app) and `description` (a textual description). The next cell shows how to access only the description field from row 200:

```python
df.loc[200, 'description']
```

Output:
```
'Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n- 5 kinds of various items and items reinforcement system\n- No network required\n- game file is as low as 20M, light-weight download!\n- supports tablet screen\n- supports Google Play Leaderboards, Achievement, Multiplay\n- supports 14 languages\n\nHomepage:\nhttps://play.google.com/store/apps/dev?id=4931745640662708567\n\nFacebook: \nhttps://www.facebook.com/spcomesgames/'
```

## Problem 1: What's in a vector?

We start by vectorising the data — more specifically, we map each app description to a tf–idf vector. This is very simple with a library like [scikit-learn](https://scikit-learn.org/stable/), which provides a [TfidfVectorizer](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html) class for exactly this purpose.  If we instantiate this class, and call `fit_transform()` on all of our app descriptions, scikit-learn will preprocess and tokenize each app description, compute tf–idf values for each of them, and return a vectorised representation:

```python
from sklearn.feature_extraction.text import TfidfVectorizer

vectorizer = TfidfVectorizer()
X = vectorizer.fit_transform(df['description'])
X
```

Output:
```
<1614x27877 sparse matrix of type '<class 'numpy.float64'>'
	with 267110 stored elements in Compressed Sparse Row format>
```

Let’s pick the app "Pancake Tower", which has a rather short description text, to see how it has been vectorised:

```python
# We can use 'toarray' to convert the sparse matrix object into a "normal" array
vec = X[1032].toarray()[0]

# The app description & its corresponding vector
df.loc[1032, 'description'], vec
```

Output:
```
("Let's see how many pancakes you can pile up!!",
 array([0., 0., 0., ..., 0., 0., 0.]))
```

That's not very informative yet.  We know that the vector contains tf–idf values, and that each dimension of the vector corresponds to a token in the vectorizer’s vocabulary; let's extract these for this specific example.

Your **first task** is to find out how to access the `vectorizer`’s vocabulary, for example by [checking the documentation of `TfidfVectorizer`](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html), and print all the tokens that are represented in the vector with a tf–idf value greater than zero (i.e., only the tokens that are actually part of this app’s description) _in descending order of the tf–idf values_.  In other words, the token with the highest tf–idf value should be at the top of your output, and the token with the lowest tf–idf value at the bottom.   Before you implement this, think about what you would expect the output look like, for example which words you would expect to have the highest/lowest tf–idf values in this example.

Your final output should look something like this:

```
<token 1>: <tf-idf value 1>
<token 2>: <tf-idf value 2>
...
```

<!-- grade_id: cell-95a3350b5c599c39, points: 0, grade: true, solution: true -->
```python
"""Print the tokens and their tf–idf values, in descending order."""

# YOUR CODE HERE
# https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html#sklearn.feature_extraction.text.TfidfVectorizer.get_feature_names_out
out_feature_names = vectorizer.get_feature_names_out()
# https://www.geeksforgeeks.org/zip-in-python/
token_tfidfval = [(token, tfidfval) for token, tfidfval in zip(out_feature_names,vec) if tfidfval>0]
#https://www.geeksforgeeks.org/sort-in-python/?ref=header_outind
token_tfidfval.sort(key=lambda x: x[1], reverse=True)
for token, tfidfval in token_tfidfval:
    print(f"{token}: {tfidfval}")
# raise NotImplementedError()
```

Output:
```
pancakes: 0.6539332651185913
pile: 0.5304701435508047
let: 0.2615287714771797
see: 0.2557630827415271
many: 0.23491959669849022
how: 0.21153246225085887
up: 0.17216837691451817
can: 0.13047602895910532
you: 0.10276923239718011
```

## Problem 2: Finding the nearest vectors

To build a small search engine, we need to be able to turn _queries_ (for example the string "pile up pancakes") into _query vectors_, and then find out which of our app description vectors are closest to the query vector.

For the first part (turning queries into query vectors), we can simply re-use the `vectorizer` that we used for the app descriptions. For the second part, an easy way to find the closest vectors is to use scikit-learn’s [NearestNeighbors](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html) class. This class needs to be _fit_ on a set of vectors (the "training set"; in our case the app descriptions) and can then be used with any vector to find its _nearest neighbors_ in the vector space.

**First,** instantiate and fit a class that returns the _ten (10)_ nearest neighbors:

<!-- grade_id: cell-bcebb6d4a62e7b7f, points: 0, grade: false, solution: true -->
```python
"""Instantiate and fit a class that returns the 10 nearest neighboring vectors."""

# YOUR CODE HERE
from sklearn.neighbors import NearestNeighbors
# https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
neighbors_nearest = NearestNeighbors(n_neighbors=10)
neighbors_nearest.fit(X)

# raise NotImplementedError()
```

Output:
```
NearestNeighbors(n_neighbors=10)
```

**Second,** implement a function that uses the vectorizer and the fitted class to find the nearest neighbours for a given query string:

<!-- grade_id: cell-18bf238a099d709b, points: 0, grade: false, solution: true -->
```python
def search(query):
    """Find the nearest neighbors in `df` for a query string.

    Arguments:
      query (str): A query string.

    Returns:
      The 10 apps (with name and description) most similar (in terms of
      cosine similarity) to the given query as a Pandas DataFrame.
    """
    # YOUR CODE HERE
    query_vector = vectorizer.transform([query])
    # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
    neigh_dist, neigh_ind = neighbors_nearest.kneighbors(query_vector)
    return(df.iloc[neigh_ind[0]][["name","description"]])
    raise NotImplementedError()
```

### 🤞 Test your code

Test your implementation by running the following cell, which will sanity-check your return value and show the 10 best search results for the query _"pile up pancakes"_:

<!-- grade_id: cell-8506d7fe5961f87b, points: 1, grade: true, solution: false -->
```python
"""Check that searching for "pile up pancakes" returns a DataFrame with ten results,
   and that the top result is "Pancake Tower"."""

result = search('pile up pancakes')
display(result)
assert isinstance(result, pd.DataFrame), "search() function should return a Pandas DataFrame"
assert len(result) == 10, "search() function should return 10 search results"
assert result.iloc[0]["name"] == "Pancake Tower", "Top search result should be 'Pancake Tower'"
success()
```

Output:
```
name  \
1032                                      Pancake Tower   
326                     Cooking School: Games for Girls   
656   Hell’s Cooking — crazy chef burger, kitchen fever   
1235                                          Solitaire   
1164                                       Rummy - Free   
1181                       Sago Mini Trucks and Diggers   
436                         Dr. Panda's Ice Cream Truck   
1442                                    Turbo Dismount™   
1446                                              UNO!™   
1326                                     TO-FU Oh!SUSHI   

                                                                                                                                                                                                                                                                                                                                                                                                                 
```

Before continuing with the next problem, play around a bit with this simple search functionality by trying out different search queries, and see if the results look like what you would expect:

```python
# Example — try out your own queries!
search("dodge trains")
```

Output:
```
name  \
1428                               Train Conductor World   
1301                                      Subway Surfers   
1300                              Subway Princess Runner   
998                       No Humanity - The Hardest Game   
228                                           Bus Rush 2   
1465                     Virus War - Space Shooting Game   
360                        Dancing Road: Color Ball Run!   
184   Bob - jigsaw puzzles free games for kids & parents   
179                       Blocky Highway: Traffic Racing   
271                            Cat Runner: Decorate Home   

                                                                                                                                                                                                                                                                                                                                                                                                       
```

## Problem 3: Custom preprocessing & tokenization

In Problem 1, you should have seen that `TfidfVectorizer` already performs some preprocessing by default and also does its own tokenization of the input data. This is great for getting started, but often we want to have more control over these steps. We can customize some aspects of the preprocessing through arguments when instantiating `TfidfVectorizer`, but for this exercise, we want to do _all_ of our preprocessing & tokenizing outside of scikit-learn.

Concretely, we want to use [spaCy](https://spacy.io), a library that we will make use of in later labs as well.  Here is a brief example of how to load and use a spaCy model:

```python
import spacy
# Load the small English model, disabling some components that we don't need right now
nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])

# Take an example sentence and print every token from it separately
doc = nlp("Apple is looking at buying U.K. startup for $1 billion")
for token in doc:
    print(token.text)
```

Output:
```
Apple
is
looking
at
buying
U.K.
startup
for
$
1
billion
```

**Your task** is to write a preprocessing function that uses spaCy to perform the following steps:
- tokenization
- lemmatization
- stop word removal
- removing tokens containing non-alphabetical characters

We recommend that you go through the [Linguistic annotations](https://spacy.io/usage/spacy-101#annotations) section of the spaCy&nbsp;101, which demonstrates how you can get the relevant kind of information via the spaCy library.

Implement your preprocessor by completing the following function:

<!-- grade_id: cell-2df4eac94cdf6be3, points: 0, grade: false, solution: true -->
```python
def preprocess(text):
    """Preprocess the given text by tokenising it, removing any stop words, 
    replacing each remaining token with its lemma (base form), and discarding 
    all lemmas that contain non-alphabetical characters.

    Arguments:
      text (str): The text to preprocess.

    Returns:
      The list of remaining lemmas after preprocessing (represented as strings).
    """
    # YOUR CODE HERE
    # https://spacy.io/usage/spacy-101#annotations
    doc = nlp(text)
    lemmas = [token.lemma_ for token in doc if token.is_stop == False and token.is_alpha==True]
    return(lemmas)
    # raise NotImplementedError()
```

### 🤞 Test your code

Test your implementation by running the following cell:

<!-- grade_id: cell-642185b139f2cef6, points: 1, grade: true, solution: false -->
```python
"""Check that the preprocessing returns the correct output for a number of test cases."""

assert (
    preprocess('Apple is looking at buying U.K. startup for $1 billion') ==
    ['Apple', 'look', 'buy', 'startup', 'billion']
)
assert (
    preprocess('"Love Story" is a country pop song written and sung by Taylor Swift.') ==
    ['Love', 'Story', 'country', 'pop', 'song', 'write', 'sing', 'Taylor', 'Swift']
)
success()
```

Output:
```
<IPython.core.display.HTML object>
```

## Problem 4: The effect of preprocessing

To make use of the new `preprocess` function from Problem 3, we need to make sure that we incorporate it into `TfidfVectorizer` and disable all preprocessing & tokenization that `TfidfVectorizer` performs by default. Afterwards, we also need to re-fit the vectorizer and the nearest-neighbors class. To make this a bit easier to handle, let’s take everything we have done so far and put it in a single class `AppSearcher`.

### Task 4.1

**Your first task** is to complete the stub of the `AppSearcher` class given below. Keep in mind:
- The `fit()` function should fit both the vectorizer (from Problem 1) and the nearest-neighbors class (from Problem 2).  Make sure to modify the call to `TfidfVectorizer` to _disable all preprocessing & tokenization_ that it would do by default, and replace it with a call to the `preprocess()` function _defined in `AppSearcher`_.
- For the `preprocess()` function, you can start by copying your solution from Problem 3.
- For the `search()` function, you can copy your solution from Problem 2.
- Make sure to adapt your code to store the everything (data, vectorizer, nearest-neighbors class) within the `AppSearcher` class, so that your solution is independent of the code you wrote above!

<!-- grade_id: cell-cb53b61e9efe98af, points: 0, grade: true, solution: true -->
```python
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        lemmas = [token.lemma_ for token in doc if token.is_stop == False and token.is_alpha==True]
        return(lemmas)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()

```

#### 🤞 Test your code

The following cell demonstrates how your class should be used. Note that it can take a bit longer to train it on the data as before, since we’re now calling spaCy for the preprocessing.

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.290758   
326                     Cooking School: Games for Girls  1.264094   
1235                                          Solitaire  1.331677   
1181                       Sago Mini Trucks and Diggers  1.339758   
1263                                   Spider Solitaire  1.345319   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.363271   
1164                                       Rummy - Free  1.372016   
436                         Dr. Panda's Ice Cream Truck  1.378334   
1245                                     Solitaire Free  1.380202   
1442                                    Turbo Dismount™  1.381175   

                                                                                                                                                                                                                                                                                                   
```

### Task 4.2

**Your second task** is to experiment with the effect of using (or not using) different preprocessing steps.  We always need to _tokenize_ the text, but other preprocessing steps are optional and require a conscious decision whether to use them or not, such as:
- lemmatization
- lowercasing all characters
- removing stop words
- removing tokens containing non-alphabetical characters

**Modify the definition of the `preprocess()` function** of `AppSearcher` to include/exclude individual preprocessing steps, run some searches, and observe if and how the results change.  Which search queries you try out is up to you — you could compare searching for "pile up pancakes" with "pancake piling", for example; or you could try entirely different search queries aimed at different kinds of apps.  (You can modify the class directly by changing the cell above under Task 4.1, or copy the definitions to the cells below, whichever you prefer; there is no separate code to show for this task, but you will use your observations here for the individual reflection.)

```python
# lemmatization
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        lemmas = [token.lemma_ for token in doc]
        return(lemmas)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()
```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.662227   
326                     Cooking School: Games for Girls  1.289065   
1235                                          Solitaire  1.347524   
1181                       Sago Mini Trucks and Diggers  1.354141   
1263                                   Spider Solitaire  1.354457   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.373068   
1164                                       Rummy - Free  1.381234   
1442                                    Turbo Dismount™  1.383333   
436                         Dr. Panda's Ice Cream Truck  1.383684   
1326                                     TO-FU Oh!SUSHI  1.386286   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.889836   
326                     Cooking School: Games for Girls  1.253859   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.360802   
1601                          من سيربح المليون الموسوعة  1.414214   
1592                               Сканворды на русском  1.414214   
1600                    كلمات وصور - كلمة السر الذكية 3  1.414214   
1212                                               Skat  1.414214   
1611                                           이사만루2019  1.414214   
1599                         دردشاتي - تعارف شات و زواج  1.414214   
742                                           Kelimelik  1.414214   

                                                                                                                                                                                                                                                                                                   
```

```python
# lowercasing all characters
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        # https://spacy.io/api/token
        lower_case = [token.text.lower() for token in doc]
        return(lower_case)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()
```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.588768   
326                     Cooking School: Games for Girls  1.295255   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.373826   
1235                                          Solitaire  1.377988   
1164                                       Rummy - Free  1.380235   
436                         Dr. Panda's Ice Cream Truck  1.382481   
1181                       Sago Mini Trucks and Diggers  1.383154   
1326                                     TO-FU Oh!SUSHI  1.384821   
1442                                    Turbo Dismount™  1.385302   
1446                                              UNO!™  1.386052   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1298                   Stylist Girl - Make Me Gorgeous!  1.379455   
1610                                              기적의 검  1.414214   
1167  RusDate - Dating. Знакомства и Общение на Русском  1.414214   
1599                         دردشاتي - تعارف شات و زواج  1.414214   
1611                                           이사만루2019  1.414214   
1595                               Филворды: поиск слов  1.414214   
1604                   ดัมมี่ - Dummy มีไพ่แคง ป๊อกเด้ง  1.414214   
686                             Hoşkin HD - İnternetsiz  1.414214   
1095                                       Pro des Mots  1.414214   
1378        Thiên Kiếm Mobile Funtap - Giang Hồ Hoàn Mỹ  1.414214   

                                                                                                                                                                                                                                                                                                   
```

```python
# removing stop words
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        rm_stop_word = [token.text for token in doc if token.is_stop == False]
        return(rm_stop_word)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()

```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.429131   
326                     Cooking School: Games for Girls  1.274510   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.368825   
1235                                          Solitaire  1.375195   
1164                                       Rummy - Free  1.380230   
436                         Dr. Panda's Ice Cream Truck  1.380312   
1442                                    Turbo Dismount™  1.383131   
427                      Dr. Panda Ice Cream Truck Free  1.387126   
1245                                     Solitaire Free  1.389436   
1326                                     TO-FU Oh!SUSHI  1.389521   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1298                    Stylist Girl - Make Me Gorgeous!  1.375465   
742                                            Kelimelik  1.414214   
1601                           من سيربح المليون الموسوعة  1.414214   
1612                                         컴투스프로야구2019  1.414214   
1592                                Сканворды на русском  1.414214   
1386  Thái Cổ Thần Vương - 3D Siêu Mượt - Tự Do Bay Lượn  1.414214   
475                                Düşen! - Kelime Oyunu  1.414214   
1610                                               기적의 검  1.414214   
1599                          دردشاتي - تعارف شات و زواج  1.414214   
1600                     كلمات وصور - كلمة السر الذكية 3  1.414214   

                                                                                                                                                                                                                                                                                         
```

```python
# removing tokens containing non-alphabetical characters
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        rm_non_alph = [token.text for token in doc if token.is_alpha==True]
        return(rm_non_alph)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()
```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.579292   
326                     Cooking School: Games for Girls  1.288873   
1235                                          Solitaire  1.364693   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.369729   
1164                                       Rummy - Free  1.371071   
436                         Dr. Panda's Ice Cream Truck  1.377365   
1442                                    Turbo Dismount™  1.380881   
1326                                     TO-FU Oh!SUSHI  1.383255   
427                      Dr. Panda Ice Cream Truck Free  1.384113   
1245                                     Solitaire Free  1.384975   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1298                    Stylist Girl - Make Me Gorgeous!  1.375352   
1167   RusDate - Dating. Знакомства и Общение на Русском  1.414214   
1612                                         컴투스프로야구2019  1.414214   
86                    Au Mobile VTC – Game nhảy Audition  1.414214   
686                              Hoşkin HD - İnternetsiz  1.414214   
1601                           من سيربح المليون الموسوعة  1.414214   
1152  Rising Super Chef - Craze Restaurant Cooking Games  1.414214   
1212                                                Skat  1.414214   
1595                                Филворды: поиск слов  1.414214   
1479                              WatchMaker Watch Faces  1.414214   

                                                                                                                                                                                                                                                                                         
```

```python
# lemmatization + lowercasing all characters
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        lemmas = [token.lemma_.lower() for token in doc]
        return(lemmas)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()
```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.632462   
326                     Cooking School: Games for Girls  1.292188   
1235                                          Solitaire  1.352310   
1181                       Sago Mini Trucks and Diggers  1.353493   
1263                                   Spider Solitaire  1.354415   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.373275   
1164                                       Rummy - Free  1.381410   
1442                                    Turbo Dismount™  1.383775   
436                         Dr. Panda's Ice Cream Truck  1.384078   
1326                                     TO-FU Oh!SUSHI  1.386151   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.872878   
326                     Cooking School: Games for Girls  1.257918   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.361076   
1610                                              기적의 검  1.414214   
475                               Düşen! - Kelime Oyunu  1.414214   
1336                                         Tallowmere  1.414214   
1609                                고스톱 PLUS (무료 맞고 게임)  1.414214   
1608                                        機動戦隊アイアンサーガ  1.414214   
1595                               Филворды: поиск слов  1.414214   
1596                                             آمیرزا  1.414214   

                                                                                                                                                                                                                                                                                                   
```

```python
# lemmatization + lowercasing all characters + removing stop words
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        lemmas = [token.lemma_.lower() for token in doc if token.is_stop == False]
        return(lemmas)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()
```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.377223   
326                     Cooking School: Games for Girls  1.275503   
1181                       Sago Mini Trucks and Diggers  1.346693   
1263                                   Spider Solitaire  1.348894   
1235                                          Solitaire  1.351876   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.368401   
1164                                       Rummy - Free  1.382413   
436                         Dr. Panda's Ice Cream Truck  1.383313   
1442                                    Turbo Dismount™  1.385583   
1160      Rope Frog Ninja Hero - Strange Gangster Vegas  1.389228   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.730324   
326                     Cooking School: Games for Girls  1.235889   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.355923   
742                                           Kelimelik  1.414214   
1600                    كلمات وصور - كلمة السر الذكية 3  1.414214   
1592                               Сканворды на русском  1.414214   
1601                          من سيربح المليون الموسوعة  1.414214   
686                             Hoşkin HD - İnternetsiz  1.414214   
844                                          Lungappen©  1.414214   
741                               Kelime Gezmece Klasik  1.414214   

                                                                                                                                                                                                                                                                                                   
```

```python
# all of them 
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # YOUR CODE HERE
        # https://stackoverflow.com/questions/41988915/avoiding-loading-spacy-data-in-each-subprocess-when-multiprocessing
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # https://www.davidsbatista.net/blog/2018/02/28/TfidfVectorizer/
        self.vectorizer = TfidfVectorizer(tokenizer=self.preprocess, preprocessor=lambda x: x, token_pattern=None)
        self.X = self.vectorizer.fit_transform(df['description'])
        self.neighbors_nearest = NearestNeighbors(n_neighbors=10)
        self.neighbors_nearest.fit(self.X)
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        # self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = self.nlp(text)
        lemmas = [token.lemma_.lower() for token in doc if token.is_stop == False and token.is_alpha==True]
        return(lemmas)
        # raise NotImplementedError()

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        query_vector = self.vectorizer.transform([query])
        # https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html
        neigh_dist, neigh_ind = self.neighbors_nearest.kneighbors(query_vector)
        name = self.df.iloc[neigh_ind[0]]['name']
        description = self.df.iloc[neigh_ind[0]]['description']
        distance = neigh_dist[0]
        return pd.DataFrame({'name': name, 'distance': distance, "description" : description})
        # raise NotImplementedError()
```

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.289759   
326                     Cooking School: Games for Girls  1.269389   
1181                       Sago Mini Trucks and Diggers  1.338611   
1235                                          Solitaire  1.340727   
1263                                   Spider Solitaire  1.345276   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.363660   
1164                                       Rummy - Free  1.372436   
436                         Dr. Panda's Ice Cream Truck  1.379078   
1442                                    Turbo Dismount™  1.381800   
1245                                     Solitaire Free  1.383728   

                                                                                                                                                                                                                                                                                                   
```

```python
apps.search("pancake piling")
```

Output:
```
name  distance  \
1032                                      Pancake Tower  0.698081   
326                     Cooking School: Games for Girls  1.227891   
656   Hell’s Cooking — crazy chef burger, kitchen fever  1.349861   
1587                    Детские песни плеер для малышей  1.414214   
1612                                        컴투스프로야구2019  1.414214   
1592                               Сканворды на русском  1.414214   
1600                    كلمات وصور - كلمة السر الذكية 3  1.414214   
1593    Слово по подсказке Ассоциации - Логические игры  1.414214   
1167  RusDate - Dating. Знакомства и Общение на Русском  1.414214   
1610                                              기적의 검  1.414214   

                                                                                                                                                                                                                                                                                                   
```

## Individual reflection

<div class="alert alert-info">
    <strong>After you have solved the lab,</strong> write a <em>brief</em> reflection (max. one A4 page) on the question(s) below.  Remember:
    <ul>
        <li>You are encouraged to discuss this part with your lab partner, but you should each write up your reflection <strong>individually</strong>.</li>
        <li><strong>Do not put your answers in the notebook</strong>; upload them in the separate submission opportunity for the reflections on Lisam.</li>
    </ul>
</div>

1. In Problem 1, which token had the highest tf–idf score, which the lowest?  Based on your knowledge of how tf–idf works, how would you explain this result?
2. Based on your observations in Problem 4, which preprocessing steps do you think are the most appropriate for this "search engine" example?  Why?

**Congratulations on finishing this lab! 👍**

<div class="alert alert-info">
    
➡️ Before you submit, **make sure the notebook can be run from start to finish** without errors.  For this, _restart the kernel_ and _run all cells_ from top to bottom. In Jupyter Notebook version 7 or higher, you can do this via "Run$\rightarrow$Restart Kernel and Run All Cells..." in the menu (or the "⏩" button in the toolbar).

</div>

```python

```